# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to understand and map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf) and [assets/NIS-PUF17-CODEBOOK.pdf](assets/NIS-PUF17-CODEBOOK.pdf). 

**Note: For all assignments please write all of your code within the function we define in order to ensure it is run by the autograder correctly**

## Question 1: proportion_of_education()
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae6, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    import pandas as pd
    import numpy as np
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)

#EDUC1 will be the variable of interest
#Where:
# 1 = less than high school
# 2 = high school
# 3 = more than high school but not college
# 4 = college degree

    df['EDUC1']
    mother_ed = np.sort(df['EDUC1'].values)
    
    prop_of_ed = {"less than high school":0,
                  "high school":0,
                  "more than high school but not college":0,
                  "college":0}

#To find the proportion, I will need total number of mothers, denoted "total_mothers"
    total_mothers = len(mother_ed)
    total_mothers

    prop_of_ed["less than high school"] = np.sum(mother_ed == 1) / total_mothers
    prop_of_ed["high school"] = np.sum(mother_ed == 2) / total_mothers
    prop_of_ed["more than high school but not college"] = np.sum(mother_ed == 3) / total_mothers
    prop_of_ed["college"] = np.sum(mother_ed == 4) / total_mothers
    
    return prop_of_ed
    
```

<!-- grade_id: cell-c0eeef201366f51d, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

Output:
```
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_57/250128802.py in <module>
----> 1 assert type(proportion_of_education())==type({}), "You must return a dictionary."
      2 assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
      3 assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
      4 assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
      5 assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

AssertionError: You must return a dictionary.
```

## Question 2: chickenpox_by_sex()
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.00779`.

<!-- grade_id: cell-a0a9e6fe67698002, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():

    import pandas as pd
    import numpy as np
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)

#     SEX == 1 denotes a MALE
#     SEX == 2 denotes a FEMALE
#
# P_NUMVRC >= 1 denotes child given one or more doses

# HAD_CPOX == 1 denotes "yes"
# HAD_CPOX == 2 denotes "no"

##############################
# Males
    males = df[df['SEX'] == 1]
    vaxed_males = males[males['P_NUMVRC'] >= 1]
    cpox_male = vaxed_males[vaxed_males['HAD_CPOX'] == 1]
    total_cpox_male = cpox_male['SEX'].count()

# males without chicken pox
    no_cpox_male = vaxed_males[vaxed_males['HAD_CPOX'] == 2]
    total_no_cpox_male = no_cpox_male['SEX'].count()

    male_ratio = total_cpox_male / total_no_cpox_male

##############################
#Similar computations for Females
    females = df[df['SEX'] == 2]
    vaxed_females = females[females['P_NUMVRC'] >= 1]
    cpox_female = vaxed_females[vaxed_females['HAD_CPOX'] == 1]
    total_cpox_female = cpox_female['SEX'].count()

# females without chicken pox
    no_cpox_female = vaxed_females[vaxed_females['HAD_CPOX'] == 2]
    total_no_cpox_female = no_cpox_female['SEX'].count()

    female_ratio = total_cpox_female / total_no_cpox_female

###############################
# Returning a dictionary with the ratios
    cpox_by_sex = {"male":male_ratio,
                  "female":female_ratio}
    
    return cpox_by_sex
    
#     raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c864, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

## Question 3: corr_chickenpox()
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no's) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses. 

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn't really the full picture, since we are not looking at when the dose was given. It's possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79b, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
                 "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})
    
    df = pd.read_csv("assets/NISPUF17.csv", index_col = 0)
    
# Indexing using 'HAD_CPOX'
# .loc attribute used to query new df
    

    df = df[df['HAD_CPOX'] < 3].loc[:,['HAD_CPOX', 'P_NUMVRC']]
    df.columns = ['had_chickenpox_column', 'num_chickenpox_vaccine_column']
#   df.isna()
    df = df.dropna()
        
    # here is some stub code to actually run the correlation
    corr, pval = stats.pearsonr(df["had_chickenpox_column"], df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    return corr

# Sanity check
#     print(corr, pval)
#     df.types()
    
#     YOUR CODE HERE
#     raise NotImplementedError()
```

<!-- grade_id: cell-73408733543a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```